# Worksheet 5: Introduction to version control


<img src="img/logo.png" width=400>

### Lecture and Tutorial Learning Goals:

After completing this week's lecture and tutorial work, you will be able to:

* Describe what version control is and why data analysis projects can benefit from it
* Create a remote version control repository on GitHub
* Move changes to files from GitHub to JupyterHub, and from JupyterHub to GitHub
* Give collaborators access to the repository
* Resolve conflicting edits made by multiple collaborators
* Communicate with collaborators using issues
* Use best practices when collaborating on a project with others

This worksheet covers parts of [Chapter 12](https://python.datasciencebook.ca/version-control.html) of the online textbook. You should read this chapter before attempting the worksheet.

**Note the following important information from UBC about your privacy: [GitHub.com](https://github.com/) is stored on servers outside Canada. When you access this site from UBC’s JupyterHub environment, you are transferred to these servers. UBC cannot guarantee security of your private information on servers outside of Canada. Please exercise caution whenever using personal information. You may wish to use a pseudonym to protect your privacy if you have concerns. Please feel free to contact us at UBC (LT.hub@ubc.ca) if you have any questions about your privacy.**

## 1. What is version control? Why use it? 

**Question 1.1** Multiple Choice:
<br> {points: 1}

Which reason listed below **is not** a good reason to use version control:

A. Version control tools provide transparency on how a project evolved by tracking the history of documents, and who made what changes to those documents.

B. Version control tools usually include a remote/cloud repository hosting service that can act as a backup of your local files (i.e., the files on your computer).

C. In practice, most data science projects involve collaboration on documents that contain code (e.g., Jupyter notebooks), and version control tools facilitate collaboration on such documents.

D. Version control tools check the accuracy of your code. 

*Assign your answer to an object called `answer1_1`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

<!-- grade_id: cell-6bdc2f57a189abcf, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-62f0d28a4d321e00, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_1)).encode("utf-8")+b"398da584ee5e21d1").hexdigest() == "11e4c24287bfefe5ab059aa602c56f6f54696cde", "type of answer1_1 is not str. answer1_1 should be an str"
assert sha1(str(len(answer1_1)).encode("utf-8")+b"398da584ee5e21d1").hexdigest() == "b717dd6dc7bdd023a7e9be9997b17edf4143dab2", "length of answer1_1 is not correct"
assert sha1(str(answer1_1.lower()).encode("utf-8")+b"398da584ee5e21d1").hexdigest() == "46c08ed6114f266b591818f2aaf8baadf0082b20", "value of answer1_1 is not correct"
assert sha1(str(answer1_1).encode("utf-8")+b"398da584ee5e21d1").hexdigest() == "4fc1ce98adafe7398515a8cff2cc005065df101f", "correct string value of answer1_1 but incorrect case of letters"

print('Success!')
```

**Question 1.2** True or false: 
<br> {points: 1}

**Git** is a remote/cloud repository hosting service where you can backup and share your files with collaborators.

*Assign your answer to an object called `answer1_2`. Make sure your answer is a boolean (e.g. `True` or `False`).* 

<!-- grade_id: cell-419826637ad2ec8b, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-decda128ce0061b5, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_2)).encode("utf-8")+b"9a10b35867588d4d").hexdigest() == "3c40317dc4b74fce459c5e2f9cd2e11fabcc8897", "type of answer1_2 is not bool. answer1_2 should be a bool"
assert sha1(str(answer1_2).encode("utf-8")+b"9a10b35867588d4d").hexdigest() == "08e8b1c5910018b83f73389ca4568d1e4c08817c", "boolean value of answer1_2 is not correct"

print('Success!')
```

## 2. Creating a space for your data science project online

For the rest of this worksheet, you will create a toy data science project on GitHub to practice using Git and GitHub. We will ask you questions about what you are doing along the way to test your understanding.

#### Signup for a free GitHub.com account:
If you do not already have a free GitHub.com account, visit [GitHub.com](https://github.com/) and signup for one. Store your username and password in a secure place (we recomend using a password manager for things like this, examples of these are [LastPass](https://www.lastpass.com/), [1Password](https://1password.com/), etc).

#### Create a GitHub repository:
On [GitHub.com](https://github.com/) create a new repository and name it `toy_ds_project`. You can decide whether to make it private or public. Ensure that you select “Add a README file.” **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#creating-a-remote-repository-on-github) in the textbook.**

**Question 2.1** Multiple Choice:

Which statement below is **not true** about GitHub repositories:
<br> {points: 1}

A. Immediately after a repository is created on GitHub.com using the website, the repository exists only on GitHub.com and does not exist on your computer (i.e., you need to do something to get a copy of it on your computer).

B. Only the creator of GitHub repository, and people the creator specify, can edit the files in the repository. This is true even when the repository is public.

C. If the repository is public, anyone on the web can view it.

D. If the repository is public, anyone on the web can edit it.

E. A GitHub repository is like a folder on Dropbox or Google Drive, but it is different in that it has special properties for version control.

*Assign your answer to an object called `answer2_1`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).* 

<!-- grade_id: cell-09da6197edcf6859, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-bcb67971fe82e848, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_1)).encode("utf-8")+b"da866d7bf00392a9").hexdigest() == "a483c51d7218ba551461533e79603f96a81c46a0", "type of answer2_1 is not str. answer2_1 should be an str"
assert sha1(str(len(answer2_1)).encode("utf-8")+b"da866d7bf00392a9").hexdigest() == "b77e1fa1c0eaa0a247494dd86ba0a8f568c847f2", "length of answer2_1 is not correct"
assert sha1(str(answer2_1.lower()).encode("utf-8")+b"da866d7bf00392a9").hexdigest() == "aa3d02e3c43e93a3b153acc3a4c9921024cd7b1b", "value of answer2_1 is not correct"
assert sha1(str(answer2_1).encode("utf-8")+b"da866d7bf00392a9").hexdigest() == "b162bfd82ae1de6e369e802c503f3a8f09acd844", "correct string value of answer2_1 but incorrect case of letters"

print('Success!')
```

## 3. Creating and editing files on GitHub

1. Edit the `README.md` file in your `toy_ds_project` repository on GitHub.com using the pen tool. Write "project creation date:" and list today's date. 
2. Commit this change directly to the main branch and write the commit message "added creation date". **This task corresponds to [this step](https://datasciencebook.ca/Getting-started-with-version-control.html#editing-files-on-github-with-the-pen-tool) in the textbook.**
3. Next, use the pen tool again to edit the `README.md` file. Write "author" and list your name as the author. Commit this change and use the commit message "added project author".
4. Explore the commit history of your project by clicking on the link that looks like this: 

<img src="img/commits_history_link.png" width=600>

> Note: you can visit the version of your repository at any stage in its history by click on the `<>` buttons! Give it a try!

**Question 3.1**  True or false:
<br> {points: 1}

Even though commit messages are required to edit a file using the pen tool on GitHub.com, it doesn't matter what message you write in practice.

*Assign your answer to an object called `answer3_1`. Make sure your answer is a boolean (e.g. `True` or `False`).* 

<!-- grade_id: cell-c7993ed9e8e63184, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-0a781125473dc3d8, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer3_1)).encode("utf-8")+b"94135b8fceb8c408").hexdigest() == "cb1854da44430b4f135a07ca39258676aa78576b", "type of answer3_1 is not bool. answer3_1 should be a bool"
assert sha1(str(answer3_1).encode("utf-8")+b"94135b8fceb8c408").hexdigest() == "9b5a99a4dcc1d11a87388f9e4ccb16821cc5c69e", "boolean value of answer3_1 is not correct"

print('Success!')
```

## 4. Cloning your repository on JupyterHub

For our data science project, we need to put a copy of our repository somewhere we can run and test the code we write (otherwise, we won't know that our code works!!!). We can use the course JupyterHub for this!

#### Clone GitHub repository to the course JupyterHub:
Clone a copy of this GitHub repository to the course JupyterHub using the Jupyter Git extension. **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#cloning-a-repository-using-jupyter) in the textbook.**


**Question 4.1** True or false:
<br> {points: 1}

The definition of **cloning a repository** is to copy/download the entire contents (files, project history, and location of the remote repository) of a remote GitHub.com repository to a computer (e.g., your workspace on a JupyterHub, or your laptop).

*Assign your answer to an object called `answer4_1`. Make sure your answer is a boolean (e.g. `True` or `False`).* 

<!-- grade_id: cell-050dd93f81d2fc95, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-f4e891dadfecbbf6, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer4_1)).encode("utf-8")+b"abdc0393a74cbfed").hexdigest() == "9f9fbb93b6cb33bc85448728b1f361726e666aa5", "type of answer4_1 is not bool. answer4_1 should be a bool"
assert sha1(str(answer4_1).encode("utf-8")+b"abdc0393a74cbfed").hexdigest() == "29add53a2e8b999484a90c60aaeb7532d7c501c9", "boolean value of answer4_1 is not correct"

print('Success!')
```

## 5. Working in a cloned repository on JupyterHub

Now that your repository exists in your workspace on the course JupyterHub, you can create a new Jupyter notebook with an Python kernel and write some code! To help this project move along, we show you below how to create a new Jupyter notebook and save it and some code to put in it. 

### Creating a new Jupyter notebook with an Python kernel

To create a new Jupyter notebook with an Python kernel in your `toy_ds_project` repository, use the file navigation menu of Jupyter so that you are inside the `toy_ds_project`:

<img src="img/to-get-to-project-root.png" width=500>


Once there, click on new Python notebook.

<img src="img/new-notebook_py.png" width=500>

Next, right-click on the filename and click on "Rename", to rename the file `marg_vs_divorce_viz.ipynb`.

<img src="img/rename.png" width=500>


### Add code to the notebook you created

Add the code below to the notebook and run it to display the data visualization. Feel free to add a narrative to the notebook if you like, commenting on the question being asked, the data visualization results, and whether correlation means causation. When you are done, save the notebook.

```python
import altair as alt
import pandas as pd


should_have_bought_butter = pd.DataFrame({
    "margarine_consumption": [8.2, 7, 6.5, 5.3, 5.2, 4, 4.6, 4.5, 4.2, 3.7],
    "maine_divorce_rate": [5, 4.7, 4.6, 4.4, 4.3, 4.1, 4.2, 4.2, 4.2, 4.1],
    "year": [2000, 2001, 2002, 2003, 2004, 2005, 2006, 2007, 2008, 2009]
})

marg_vs_time = alt.Chart(should_have_bought_butter).mark_line(color="steelblue").encode(
    x=alt.X("year").title("Year"),
    y=alt.Y("margarine_consumption")
        .title("Margarine consumption (lbs per capita)")
        .scale(zero=False)
).properties(
    height=250
)


divorce_rate_vs_time = marg_vs_time.mark_line(color="coral").encode(
    y=alt.Y("maine_divorce_rate")
        .title("Divorce rate in Maine (per 1000")
        .scale(zero=False)
)

(marg_vs_time | divorce_rate_vs_time).properties(
    title="Divorce rate in Maine correlates with margarine consumption"
)
```

## 6. Specifying files to commit

Now we would like to start the process of putting `marg_vs_divorce_viz.ipynb` under version control and eventually push this file to our remote repository on GitHub.com. The first step to doing this is to **add** the changes to this file (creating it and the code) to the Git staging area. Go ahead and use the Jupyter Git extension to do this now. **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#specifying-files-to-commit) in the textbook.**

**Question 6.1** Multiple Choice:
<br> {points: 1}

Git has a distinct step of **adding** files to the staging area because:

A. Not all changes we make (i.e., files we create or edit) are ones that we want to push to our remote GitHub repository.

B. It allows us to edit multiple files at once, but associate particular commit messages with particular files (so that the commit messages can more specifically reflect the changes that were made).

C. This is technically required of all version control software.

D. A and C.

E. A and B.

*Assign your answer to an object called `answer6_1`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).* 

<!-- grade_id: cell-78a41a64da6293e8, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-4df7797ab4b028cf, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer6_1)).encode("utf-8")+b"fd6b9e20cb45ad46").hexdigest() == "e1d4ceebb49be8ab86859e0dc89ec2174459ab47", "type of answer6_1 is not str. answer6_1 should be an str"
assert sha1(str(len(answer6_1)).encode("utf-8")+b"fd6b9e20cb45ad46").hexdigest() == "e44cdcea301a5301b1699729dbbbb25da10ade42", "length of answer6_1 is not correct"
assert sha1(str(answer6_1.lower()).encode("utf-8")+b"fd6b9e20cb45ad46").hexdigest() == "f1b9544e121af724a484c17323be827340ca4b71", "value of answer6_1 is not correct"
assert sha1(str(answer6_1).encode("utf-8")+b"fd6b9e20cb45ad46").hexdigest() == "f94b7f752871246b502709f58766f5f6a60f7313", "correct string value of answer6_1 but incorrect case of letters"

print('Success!')
```

## 7. Making the commit 

The next step is to **commit** our changes to our local Git repository. You can use the Jupyter Git extension to do this now. **This task corresponds to [this step](https://datasciencebook.ca/Getting-started-with-version-control.html#making-the-commit) in the textbook.**

**Question 7.1** True or false:
<br> {points: 1}

When we **commit** our changes to Git, the snapshot of changes, the commit message, the time and date stamp, and the user who committed the changes are all saved to the Git history on GitHub.

*Assign your answer to an object called `answer7_1`. Make sure your answer is a boolean (e.g. `True` or `False`).* 

<!-- grade_id: cell-abcee8a2afc5bc19, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-a4297f9a9b8fe5a9, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer7_1)).encode("utf-8")+b"0edda0f4b18aac5b").hexdigest() == "63a7d020f048a72ab59dcde97fdf4667ea9a0470", "type of answer7_1 is not bool. answer7_1 should be a bool"
assert sha1(str(answer7_1).encode("utf-8")+b"0edda0f4b18aac5b").hexdigest() == "2a892d14da42ebfa1a44b819095019334c1f8dd6", "boolean value of answer7_1 is not correct"

print('Success!')
```

## 8. Pushing the commits to GitHub

Finally, we are ready to send our changes (creating and adding code to `marg_vs_divorce_viz.ipynb`) to our remote repository through a process we call "pushing". Go ahead and do this now. **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#pushing-the-commits-to-github) in the textbook.**

After completing pushing your work to the remote repository on GitHub, visit your repository on GitHub.com and check out what your awesome toy project looks like!!!

**Question 8.1** Multiple Choice:

Which statement below is **not** true?
<br> {points: 1}

A. Cloning and pulling a GitHub repository are the exact same thing.

B. Pushing with Git is the act of sending changes that were committed to Git to a remote repository, for example, on GitHub.com.

C. Pulling with Git is the act of collecting changes that exists in a remote repository, for example, on GitHub.com, that do not yet exist on the local computer you are working on (i.e., your workspace on the JupyterHub or your laptop).

D. You should push your work to GitHub anytime you want to share your work with others, or when you are done a work session and want to back up your work.

*Assign your answer to an object called `answer8_1`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).* 

<!-- grade_id: cell-513a65b2e4fa6c01, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-ac98c61c35f6f0b6, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer8_1)).encode("utf-8")+b"1ed95697df6bd6fd").hexdigest() == "324f98812e5bd88792e46d45a1928677bd6b6d67", "type of answer8_1 is not str. answer8_1 should be an str"
assert sha1(str(len(answer8_1)).encode("utf-8")+b"1ed95697df6bd6fd").hexdigest() == "6a9d5b89022e5d75293cef63dd3b59d85aa3124a", "length of answer8_1 is not correct"
assert sha1(str(answer8_1.lower()).encode("utf-8")+b"1ed95697df6bd6fd").hexdigest() == "477e6eccb34bc2dd87ea2de2b1e26cec32494b64", "value of answer8_1 is not correct"
assert sha1(str(answer8_1).encode("utf-8")+b"1ed95697df6bd6fd").hexdigest() == "26c4801a2701d317142596fbb653b77bf7fcbbdd", "correct string value of answer8_1 but incorrect case of letters"

print('Success!')
```

## 9. Giving collaborators access to your project

One of the advantages of using version control tools, such as Git and GitHub, is how it lets you collaborate. Let's get some practice starting down this path. Add one or more of your breakout group members to your GitHub repository as a collaborator. **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#giving-collaborators-access-to-your-project) in the textbook.**

**Question 9.1** True or false:
<br> {points: 1}

You can clone or pull from any public remote repository on GitHub.com, however you can only push to public remote repositories on GitHub.com that you own are a collaborator on.

*Assign your answer to an object called `answer9_1`. Make sure your answer is a boolean (e.g. `True` or `False`).* 

<!-- grade_id: cell-e27a72ed06d0a5c6, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-9dc8d28efb9ccda7, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer9_1)).encode("utf-8")+b"7d9bc89c71fafc21").hexdigest() == "2f111c707976b8afa6f40a4669cfe079aead440e", "type of answer9_1 is not bool. answer9_1 should be a bool"
assert sha1(str(answer9_1).encode("utf-8")+b"7d9bc89c71fafc21").hexdigest() == "d0d6d1223c081f806119f49af0cdbea463e1ef0d", "boolean value of answer9_1 is not correct"

print('Success!')
```

#### (Optional) more collaboration practice!

If you want to practice more Git & GitHub skills for collaboration, ask someone in your team if you can collaborate and send an edit to their project. To do this, they will need to add you as a collaborator, and then you will need to clone their repository to your JupyterHub. After that, you can edit some files (or create a whole new one), save your work, and then use the Jupyter Git extension to add, commit, and push your changes to their remote GitHub repository.

## 10. Communicating using GitHub issues

It's easy for project communications to get lost in email or whatever messaging platform you use to communicate with your team. GitHub issues are an excellent tool explicitly designed for project collaboration as they are "attached" to the project's remote GitHub repository. Your task here is to go to the issue tab for your project and create an issue about something you might want to improve about your project. **This task corresponds to [this step](https://python.datasciencebook.ca/version-control.html#communicating-using-github-issues) in the textbook.**

**Question 10.1** Multiple Choice:
<br> {points: 1}

Which statement below is **not** a reason why GitHub issues are an ideal medium for project-specific communications?

A. Issues are part of each GitHub repository, and thus "attached" to the project.

B. Issues only persist while they are open, and immediately deleted when they are closed.

C. Issues are easily searchable using GitHub’s search tools.

D. All issues are accessible to all project collaborators, so no one is left out of the conversation.

E. Issues can be set up so that team members get email notifications when a new issue is created or a new post is made in an issue thread.

*Assign your answer to an object called `answer10_1`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).* 

<!-- grade_id: cell-79a4febe0e471e05, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-6656677ba896c92f, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer10_1)).encode("utf-8")+b"5af9632f50d95e06").hexdigest() == "23aec8bf244ba441b42a2ffa449d9eb4e5aa9dba", "type of answer10_1 is not str. answer10_1 should be an str"
assert sha1(str(len(answer10_1)).encode("utf-8")+b"5af9632f50d95e06").hexdigest() == "33538b8a4701890577f526542ccd7f3fef4e8910", "length of answer10_1 is not correct"
assert sha1(str(answer10_1.lower()).encode("utf-8")+b"5af9632f50d95e06").hexdigest() == "0cf92f8d3ad6283f7d8c7845f3d54b7bef4efac3", "value of answer10_1 is not correct"
assert sha1(str(answer10_1).encode("utf-8")+b"5af9632f50d95e06").hexdigest() == "059247f5c812f8b412d2b832b34f184287410cc3", "correct string value of answer10_1 but incorrect case of letters"

print('Success!')
```

#### (Optional) Even more collaboration practice! 

Visit a team member's GitHub repository and leave a polite but constructive message on how they could improve their project. 


## Nice work! You're all done here!

<img src="https://media.giphy.com/media/lcYFNTaz4U9jy/giphy.gif">